# SLU07 - Function Fundamentals: Exercise notebook

Let's practice some functions! The important part here is that you can take some plain old english and translate it into a function.

```python
# Don't mind this here, just some imports for the exercises - MAKE SURE YOU RUN THIS CELL
import inspect
from io import StringIO 
import sys

class Capturing(list):
    def __enter__(self):
        self._stdout = sys.stdout
        sys.stdout = self._stringio = StringIO()
        return self
    def __exit__(self, *args):
        self.extend(self._stringio.getvalue().splitlines())
        del self._stringio    # free up some memory
        sys.stdout = self._stdout
```

### Exercise 1

Let's start with a simple one!

Implement a function that:
- Is named `fun_adding_numbers`
- Takes two arguments
- Adds the arguments together
- Returns the value of the numbers added together


<!-- grade_id: cell-a3913c46bc1f65f2, points: 0, grade: false, solution: true -->
```python
# def fun_adding_numbers...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b7d97d9eb6d71625, points: 2.5, grade: true, solution: false -->
```python
sig = inspect.signature(fun_adding_numbers)

assert fun_adding_numbers(1, 2) == 3
assert fun_adding_numbers(4, 5) == 9
assert '+' in inspect.getsource(fun_adding_numbers)
assert len(sig.parameters) == 2
```

### First function done!! It's as simple as that! 😎 

### Exercise 2

Let's go for another one! 

Implement a function that

- Is named `fun_multiplying_numbers`
- Takes two parameters
- Multiplies the arguments together
- Returns the value of the numbers multiplied together


<!-- grade_id: cell-b5cba4d5bddbb4b5, points: 0, grade: false, solution: true -->
```python
# def fun_multiplying_numbers...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-5db0d0c172007a90, points: 2.5, grade: true, solution: false -->
```python
source = inspect.signature(fun_multiplying_numbers)

assert fun_multiplying_numbers(1, 2) == 2
assert fun_multiplying_numbers(4, 5) == 20
assert '*' in inspect.getsource(fun_multiplying_numbers)
assert len(source.parameters) == 2
```

### Exercise 3

Now for some special functions points! You need to think about this one and what you have learned up until now.

Try to implement the previous `fun_multiplying_numbers` function but now with a twist:
- Is named `fun_multiplying_numbers_with_a_twist`
- Takes two parameters
- **Multiplies the arguments without using the * operator!**
- Returns the value of the multiplication

Tip: If you remember the good old days of math, multiplication is just a sequence of adding operations in a loop and you already learned how to do those!

<!-- grade_id: cell-24f09cf30d150f7a, points: 0, grade: false, solution: true -->
```python
# def fun_multiplying_numbers_with_a_twist(a,b):
    # result = 0
    # (loop)
    # return result

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9038267e2c8c00c4, points: 2.5, grade: true, solution: false -->
```python
source = inspect.signature(fun_multiplying_numbers_with_a_twist)

assert fun_multiplying_numbers_with_a_twist(1, 2) == 2
assert fun_multiplying_numbers_with_a_twist(4, 5) == 20
assert fun_multiplying_numbers_with_a_twist(100, 200) == 20000
assert '*' not in inspect.getsource(fun_multiplying_numbers_with_a_twist), "Don't use the multiplication operator! 😉"
assert len(source.parameters) == 2
```

### Exercise 4

Implement a function that

- Is named `concat_strings`
- Takes two parameters that are strings
- Concatenates the strings together in the order that they are received
- Returns the concatenated strings
- **Twist**: the second parameter assumes as a default string " best" -> *mind the initial space!*

<!-- grade_id: cell-90e04d850de62687, points: 0, grade: false, solution: true -->
```python
#Remember the default values 🤔 

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d309394f7d502266, points: 2.5, grade: true, solution: false -->
```python
source = inspect.getsource(concat_strings)

sig = inspect.signature(concat_strings)

assert concat_strings('hello', ' world') == 'hello world'
assert concat_strings('Daniel', ' will give me a good grade 🙏') == 'Daniel will give me a good grade 🙏'
assert concat_strings('This course is the') == 'This course is the best'
assert concat_strings('functions are the') == 'functions are the best'
assert len(sig.parameters) == 2
```

### Exercise 5

Implement a function that

- Is named `convert_to_int_and_add`
- Takes two parameters
- Both of the parameters are integers **represented as strings** (for example, "9")
- Converts the strings to integers
- Add the integers together
- Return the value of the integers added together

<!-- grade_id: cell-5769f897199017e3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-bf3c5cd42c49fa00, points: 2.5, grade: true, solution: false -->
```python
source = inspect.getsource(convert_to_int_and_add)

sig = inspect.signature(convert_to_int_and_add)

assert convert_to_int_and_add('1', '2') == 3
assert convert_to_int_and_add('5', '100') == 105
assert '+' in source
assert 'int' in source
assert len(sig.parameters) == 2
```

### Exercise 6

Implement a function that

- Is named `two_step`
- Takes 4 parameters
- Multiplies the first 2 together and stores the output in a variable called `step1`
- Divides 3rd parameter by the 4th parameter and stores the output in a variable called `step2`
- Adds the variables `step1` and `step2` together, calling it `result`
- If `result` is greater than 20, subtract 10 to it
- Returns the final result

<!-- grade_id: cell-70188605f9305dc8, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9712fb50a8f3d533, points: 2.5, grade: true, solution: false -->
```python
source = inspect.getsource(two_step)

sig = inspect.signature(two_step)

assert two_step(2, 2, 4, 2) == 6.0, "The final result is not matching"
assert two_step(2, 3, 100, 5) == 16.0, "Are you subtracting correctly? 🤔"
assert '*' in source
assert '/' in source
assert 'step1' in source
assert 'step2' in source
assert 'result' in source
assert len(sig.parameters) == 4
```

### Exercise 7

Write a function that

1. Is named `holla`
1. Takes 5 arguments
1. Calls `fun_adding_numbers` with the first two arguments and stores the output in `out1`
1. Calls `fun_multiplying_numbers` with the second two arguments and stores the output in `out2`
1. Calls `two_step` with the last 4 arguments and stores the output in `out3`
1. Returns the addition of `out1`, `out2`, and `out3`

<!-- grade_id: cell-a5e8cf3762c2538a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a02c4760a2ce92bf, points: 2.5, grade: true, solution: false -->
```python
source = inspect.getsource(holla)
sig = inspect.signature(holla)

assert holla(1, 2, 2, 4, 2) == 17
assert holla(1, 2, 3, 100, 5) == 319.0
assert 'fun_adding_numbers(' in source
assert 'fun_multiplying_numbers(' in source
assert 'two_step(' in source
assert 'out1 + out2 + out3' in source
assert len(sig.parameters) == 5
```

### Exercise 8

Write a function that

1. Is called `add_and_print`
1. Takes two arguments
1. Adds them together
1. Prints the output
1. Returns the value `5`

<!-- grade_id: cell-5a9e41f05baeb9b5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-733a771bc8560db2, points: 2.5, grade: true, solution: false -->
```python
source = inspect.getsource(add_and_print)
sig = inspect.signature(add_and_print)

with Capturing() as output:
    retval = add_and_print(1, 2)

assert retval == 5
assert output[0] == '3'

with Capturing() as output:
    retval = add_and_print(10287366, 912865763524)
    
assert retval == 5
assert output[0] == '912876050890'

assert '+' in source
assert len(sig.parameters) == 2
```

# Submit your work!

To submit your work, [follow the steps here, in the section "Grading the Exercise Notebook"!](https://github.com/LDSSA/ds-prep-course-2022#22---working-on-the-learning-units)