# SLU 17 - Exam Prep I

## Batch 1 Python Exam 
This is the Python exam from batch 1. Exam Duration: 2h

```python
import json
#used for evaluation
import inspect
import math
import re
import json
import sys
```

## Question 1

Complete the function `import_data`. This function should:
- Open the **countries.json** file that is located within the folder name **data** ;
- Use Python’s module json to load the data from the file into a dictionary;
- Return the dictionary;

<!-- grade_id: cell-ce9179c0612fd6b0, points: 0, grade: false, solution: true -->
```python
def import_data():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-78d54c806fb28266, points: 2, grade: true, solution: false -->
```python
countries = import_data()
assert isinstance(countries, dict), "Not correct. Keep trying."
assert len(countries) == 194, "Not correct. Keep trying."
assert countries["Albania"]["Continent"] == "Europe", "Not correct. Keep trying."
assert countries["Vietnam"]["Area"] == 329560, "Not correct. Keep trying."
print("---- All asserts passed ---- ")
```

## Question 2

The dictionary that you built contains information about a serie of countries. So from now on, we’ll call it the `countries` dictionary.
In the `countries` dictionary, each country has a field called `Population` (be aware that this is case sensitive) that gives the number of people living in that country. Since this is a count, the data should have integer type. However, we have found some numbers with type of *float*. Let's try to fix this!

Complete the function `convert_population_to_int`. This function should:
- Receive the `countries` dictionary as argument;
- Convert the `Population` field from float to int type for all the countries, using a `for` loop;
- Return the updated `countries` dictionary;

<!-- grade_id: cell-28d89deeef261406, points: 0, grade: false, solution: true -->
```python
def convert_population_to_int(countries):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-33b501c79f675ab0, points: 2, grade: true, solution: false -->
```python
countries = convert_population_to_int(countries)
for country in countries.keys():
    assert isinstance(countries[country]["Population"], int),  "Not correct. Keep trying."
assert len(countries) == 194, "Not correct. Keep trying."
source = inspect.getsource(convert_population_to_int)
assert "for" in source, "Not correct. Keep trying."
print("---- All asserts passed ---- ")
```

## Question 3

In the countries dictionary, each country has a field called `Area`. This is the area of the country given in square miles.

Complete the function `convert_area_to_sq_km`. This function should:
- Receive the `countries` dictionary as argument;
- Convert the `Area` field from square miles to square kilometres;
 - Use the following conversion: 1 sq mi = 2.58999 sq km;
 - Round the result to 1 decimal digit;
- Return the updated `countries` dictionary;

<!-- grade_id: cell-3fd27539b41bed5c, points: 0, grade: false, solution: true -->
```python
def convert_area_to_sq_km(countries):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-b77708fde98d925a, points: 2, grade: true, solution: false -->
```python
countries = convert_area_to_sq_km(countries)
for country in countries.keys():
    assert isinstance(countries[country]["Area"], float), \
            "Not correct. Keep trying."
assert len(countries) == 194, "Not correct. Keep trying."
assert math.isclose(countries["Haiti"]["Area"], 71872.2, abs_tol=0.1), \
       "Not correct. Keep trying."
assert math.isclose(countries["Morocco"]["Area"], 1156560.0, abs_tol=0.1), \
       "Not correct. Keep trying."
print("---- All asserts passed ---- ")
```

## Question 4

Complete the function `get_europe_countries`. This function should:
- Receive the `countries` dictionary as argument;
- Build a list with the names of the countries in Europe sorted by alphabetical order (from A to Z);
 - Use the field called `Continent`;
 - Use list comprehension;
 - Use the sort method from the list data type;
- Return the list;

<!-- grade_id: cell-ee1f38ef730c2004, points: 0, grade: false, solution: true -->
```python
def get_europe_countries(countries):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-48bf94a05b056aa0, points: 2, grade: true, solution: false -->
```python
countries_list = get_europe_countries(countries)
assert isinstance(countries_list, list), "Not correct. Keep trying."
assert len(countries_list) == 35, "Not correct. Keep trying."
assert countries_list[0] == "Albania", "Not correct. Keep trying."
assert countries_list[5] == "Croatia", "Not correct. Keep trying."
assert countries_list[34] == "United Kingdom", "Not correct. Keep trying."
source = inspect.getsource(get_europe_countries)
assert "for" in source, "Not correct. Keep trying."
assert "[" in source, "Not correct. Keep trying."
assert "in" in source, "Not correct. Keep trying."
assert ".sort(" in source, "Not correct. Keep trying."
print("---- All asserts passed ---- ")
```

## Question 5

Complete the function `get_literacy_levels_by_continent`. This function should:
- Receive the `countries` dictionary as argument;
- Receive a continent name as argument;
- For the countries in the continent (received as argument), compute the literacy level as:
 - literacy in [0, 25[ % - VERY_LOW;
 - literacy in [25, 50[ % - LOW;
 - literacy in [50, 70[ % - MEDIUM;
 - literacy in [70, 90[ % - HIGH;
 - literacy in [90, 100] % - VERY_HIGH;
- Return a list of tuples like: [(country_1, literacy_1, literacy_level_1), ..., (country_n, literacy_n, literacy_level_n), ...];

<!-- grade_id: cell-3911f29ed53cd968, points: 0, grade: false, solution: true -->
```python
def get_literacy_levels_by_continent(countries, continent):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-fa41a51f9ff00d7c, points: 2, grade: true, solution: false -->
```python
for continent in ["Europe","Africa"]:
    literacy_tuple = get_literacy_levels_by_continent(countries, continent)    
    for literacy in literacy_tuple:
        assert len(literacy) == 3, "Not correct. Keep trying."
        assert isinstance(literacy[0], str), "Not correct. Keep trying."
        assert isinstance(literacy[1], float), "Not correct. Keep trying."
        assert isinstance(literacy[2], str), "Not correct. Keep trying."

    if continent == "Europe":
        assert len(literacy_tuple) == 35, "Not correct. Keep trying."
        assert literacy_tuple[21][0] == 'Malta', "Not correct. Keep trying."
        assert math.isclose(literacy_tuple[21][1], 92.8, abs_tol=0.1), \
               "Not correct. Keep trying."
        assert literacy_tuple[21][2] == 'VERY_HIGH', "Not correct. Keep trying."
        assert literacy_tuple[1][2] == 'VERY_HIGH', "Not correct. Keep trying."
        assert literacy_tuple[0][2] == 'HIGH', "Not correct. Keep trying."
        assert literacy_tuple[-1][2] == 'VERY_HIGH', "Not correct. Keep trying."

    if continent == "Africa":
        assert len(literacy_tuple) == 55, "Not correct. Keep trying."
        assert literacy_tuple[14][0] == 'Djibouti', "Not correct. Keep trying."
        assert math.isclose(literacy_tuple[14][1], 67.9, abs_tol=0.1), \
               "Not correct. Keep trying."
        assert literacy_tuple[14][2] == 'MEDIUM', "Not correct. Keep trying."
        assert literacy_tuple[1][2] == 'LOW', "Not correct. Keep trying."
        assert literacy_tuple[0][2] == 'HIGH', "Not correct. Keep trying."
        assert literacy_tuple[-1][2] == 'VERY_HIGH', "Not correct. Keep trying."
print("---- All asserts passed ---- ")
```

## Question 6

Complete the function `get_country_codes`. This function should:
- Receive the `countries` dictionary as argument;
- Build a list with all the countries names;
- Using `map` and a `lambda` function, convert the country names in country codes by selecting the 3 first letters of the country;
- Returns this list;

Note: Don't worry about duplicates on the result.

<!-- grade_id: cell-bd068242631b4909, points: 0, grade: false, solution: true -->
```python
def get_country_codes(countries):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-5f1eaf642616c8a1, points: 2, grade: true, solution: false -->
```python
country_codes = get_country_codes(countries) 
assert len(country_codes) == 194, "Not correct. Keep trying."
for code in country_codes:
    assert isinstance(code, str), "Not correct. Keep trying."
    assert len(code) == 3, "Not correct. Keep trying."
assert country_codes[134] == "Pak", "Not correct. Keep trying."
assert country_codes[62] == "Fra", "Not correct. Keep trying."

source = inspect.getsource(get_country_codes)
assert "map" in source, "Not correct. Keep trying."
assert "lambda" in source, "Not correct. Keep trying."
print("---- All asserts passed ---- ")
```

## Question 7

The `Country` class is going to be used to represent countries and to return information about them.

### Question 7.1

Complete the `get_population_in_millions` method. This method should:
- Return the population of the country that the class represents;
- The population should be returned in millions with 2 decimal digits;

<!-- grade_id: cell-9b9ae4ba8e61e508, points: 0, grade: false, solution: true -->
```python
class Country:
    def __init__(self, country_name, population, continent):
        self.country_name = country_name
        self.population = population
        self.continent = continent

    def get_population_in_millions(self):
        # YOUR CODE HERE
        raise NotImplementedError()
```

<!-- grade_id: cell-9e28e403f317f23d, points: 2, grade: true, solution: false -->
```python
country = Country("Penguinea", 12345678, "Antarctica")
assert isinstance(country, Country), "Not correct. Keep trying."
assert isinstance(country.get_population_in_millions(), float), "Not correct. Keep trying."
assert math.isclose(country.get_population_in_millions(), 12.35, abs_tol=0.01),"Not correct. Keep trying."

country = Country("Walrussia", 8344567, "Antarctica") 
assert math.isclose(country.get_population_in_millions(), 8.34, abs_tol=0.01), "Not correct. Keep trying."
print("---- All asserts passed ---- ")
```

## Question 7.2

Complete the `get_country_population`. This function should:
- Receive the `countries` dictionary as argument;
- Receive a country name as argument;
- Try to find the country (received as argument) in the countries dictionary;
- If the country doesn’t exist in the countries dictionary, catch the KeyError, print a statement with the information that there is no information for that country and return `None`;
- If the country exists, create an object of the `Country` class;
- Call the `get_population_in_millions` method for the country;
- Return the population in millions value;

<!-- grade_id: cell-803b263dbee2e300, points: 0, grade: false, solution: true -->
```python
def get_country_population(countries, country_name):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-0d0db73f5256c8e5, points: 6, grade: true, solution: false -->
```python
assert math.isclose(get_country_population(countries, "Turkey"), 70.41, abs_tol=0.01), "Not correct. Keep trying."
assert get_country_population(countries, "Narnia") is None,  "Not correct. Keep trying."
print("---- All asserts passed ---- ")
```

### Last but not least, get your grade!
To know your grade, run the following command in the **COMMAND LINE**:

`python grading.py`